rule_files:
  - planner-alert-rules.yml

evaluation_interval: 1m

tests:
  - interval: 1m
    input_series:
      - series: 'agent_planner_route_total{env="prod"}'
        values: '0+100x30'
      - series: 'agent_planner_fallback_total{env="prod"}'
        values: '0+10x30'
    alert_rule_test:
      - eval_time: 20m
        alertname: PlannerFallbackSpikeWarningProd
        exp_alerts: []
      - eval_time: 20m
        alertname: PlannerFallbackRatioCriticalProd
        exp_alerts: []

  - interval: 1m
    input_series:
      - series: 'agent_planner_route_total{env="prod"}'
        values: '0+10x30'
      - series: 'agent_planner_fallback_total{env="prod"}'
        values: '0+8x30'
    alert_rule_test:
      - eval_time: 20m
        alertname: PlannerFallbackRatioCriticalProd
        exp_alerts:
          - exp_labels:
              alertname: PlannerFallbackRatioCriticalProd
              service: agent-app
              module: planner
              severity: critical
              env: prod

  - interval: 1m
    input_series:
      - series: 'agent_planner_route_total{env="staging"}'
        values: '0+1x30'
      - series: 'agent_planner_fallback_total{env="staging"}'
        values: '0+11x30'
    alert_rule_test:
      - eval_time: 20m
        alertname: PlannerFallbackSpikeWarningStaging
        exp_alerts:
          - exp_labels:
              alertname: PlannerFallbackSpikeWarningStaging
              service: agent-app
              module: planner
              severity: warning
              env: staging
